# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Primeiro e último valor da lista

Implemente uma função que retorne uma lista com com primeiro valor e o último valor da lista recebida como parâmetro:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def start_end_list(values):
    """
    Retorna uma lista com o primeiro e último valor da lista.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de start_end_list()"

assert start_end_list([7, 5, 21, 12, 566, 65]) == [7, 65]
assert start_end_list([3, 65, 3, 64, 3, 643, 0]) == [3,0]
assert start_end_list([61, 16, 34, 34, 43, 1, 6]) == [61, 6]
assert start_end_list([852]) == [852, 852]
print('Parabéns!')
```