# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Sistema de Abastacimento de Carros

A classe Carro mostrada abaixo possui dois atributos que representam, respectivamente, o consumo de combustível (medido em km/litro) e a quantidade de combustível no tanque. O consumo é especificado no construtor e o nível de combustível inicial é 0.

Forneça um método drive( ) que simule o ato de dirigir o veículo por uma certa distância, reduzindo o nível de combustível no tanque de gasolina. Caso o nível for ficar negativo, deve se manter como Zero.

Forneça um método showGas( ), que retorna o nível atual de combustível.

Forneça um método addGas( ), para abastecer o tanque.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
class Car:
    def __init__(self, kilometersPerLiter):
        self.kilometersPerLiter = kilometersPerLiter;
        self.gasQuantity = 0;
    
    # Ensina o Python como testar se duas pessoas são iguais. Não modifique.
    def __eq__(self, other):
        return self.kilometersPerLiter == other.kilometersPerLiter and self.gasQuantity == other.gasQuantity;
    
    def drive(self, distance):
        """
        Simula o ato de dirigir o veículo por uma certa distância, reduzindo o nível de combustível.
        """
        NotImplemented # YOUR ANSWER HERE

    def showGas(self):
        """
        Retorna o nível atual de combustível.
        """
        NotImplemented # YOUR ANSWER HERE

    def addGas(self, gasAdded):
        """
        Abastece o tanque.
        """
        NotImplemented # YOUR ANSWER HERE
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
# Exemplo para você experimentar sua implementação :)

myCar = Car(15);                    # 10 quilômetros por litro de combustível.
"{}".format(myCar.showGas());       # Imprime o combustível atual no tanque.
myCar.drive(60);                    # anda 60 quilômetros.
myCar.addGas(20);                   # abastece com 20 litros de combustível. 
myCar.drive(30);                    # anda 30 quilômetros.
"{}".format(myCar.showGas());       # Imprime o combustível que resta no tanque.
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
# Testes do Sistema

theCar = Car(10);
assert theCar.showGas() == 0;
myCar.addGas(20);
assert theCar.showGas() == 20;
myCar.drive(50);
assert theCar.showGas() == 15;
myCar.drive(100);
assert theCar.showGas() == 5;
myCar.addGas(95);
assert theCar.showGas() == 100;
myCar.drive(900);
assert theCar.showGas() == 10;
myCar.drive(110);
assert theCar.showGas() == 0;
print('Parabens');
```