# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Divisibilidade

Implemente uma função que informe se o primeiro valor inserido é divisível pelo segundo valor:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def is_divisible(number_1, number_2):
    """
    Retorna True se for divisível e False se não for
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-recuperar-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função is_divisible()"

assert is_divisible(10, 2) == True
assert is_divisible(10, 3) == False
assert is_divisible(-3, -3) == True
print('Parabéns!')
```