# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Conversão em segundos

Escreva uma função que convertará dias, horas ou minutos em segundos, ela receberá 2 parâmetros e retornará um valor inteiro. O primeiro parâmetro poderá ser de 3 tipos: "DIA", "HORA" ou "MINUTO", o segundo será o valor correspondente ao primeiro parâmetro. Por exemplo, deseja-se converter 3 dias em segundos, deve-se passar como primeiro parâmetros para a função o nome "DIA" e como segundo parâmetros o valor de dias, que nesse caso é 3. Caso um valor negativo seja introduzido, a função deve retornar 0.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def secondsConverter(unidadeDeTempo, valor):
    """
    Retorna o valor em segundos
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de secondsConverter()"

assert secondsConverter('DIA', 3) == 259200
assert secondsConverter('DIA', 30) == 864000
assert secondsConverter('DIA', -9) == 0
assert secondsConverter('HORA', 2) == 7200
assert secondsConverter('HORA', 17) == 61200
assert secondsConverter('HORA', -100) == 0
assert secondsConverter('MINUTO', 1) == 60
assert secondsConverter('MINUTO', 560) == 33600
assert secondsConverter('MINUTO', 0) == 0
print('Parabéns!')
```