# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Contador de Bits

Escreva uma função que receba um inteiro como entrada e retorne o número de bits que são iguais a 1 na representação binária desse número. Você pode assumir que a entrada não será negativa.

Exemplo: A representação binária de 1234 é 10011010010, então a função deve retornar 5 neste caso.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def bit_counting(value):
    """
    Retorna um inteiro com a contagem de números 1's
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
# Exemplo para você experimentar sua implementação :)
print("{}".format(bit_counting(1234)));       # Imprime o valor da função
```

Output:
```
5
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
# Testes do Sistema

assert bit_counting(0) == 0;
assert bit_counting(1) == 1;
assert bit_counting(2) == 1;
assert bit_counting(22) == 3;
assert bit_counting(8191) == 13;
assert bit_counting(82) == 3;
assert bit_counting(876321054) == 16;
assert bit_counting(123) == 6;
assert bit_counting(745896321) == 15;
assert bit_counting(1234) == 5;

print('Parabens');
```

Output:
```
Parabens
```