# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

This question was adapted from this: https://github.com/fabiommendes/desenvolvimento-de-software/blob/master/provas/teste-2.ipynb.
And its use has been authorized by its repository owner.

---
## Sistema de registro

Um sistema possui um módulo de autenticação em que os usuários podem iniciar uma seção após a verificação de senha e e-mail. As informações dos usuários são armazenadas na lista USUARIOS. Você deve implementar as funções de login() e recuperar_senha() para passar nos testes abaixo.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
# Não edite a lista de usuários
USUARIOS = [
    {
        'nome': 'Ada Lovelace',
        'emails': ['ada@gmail.co.uk', 'algo-ada@hotmail.com'],
        'nascimento': (1815, 12, 10),
        'senha': 'algoritmo',
    },
    {
        'nome': 'Alan Turing',
        'emails': ['turing@gmail.co.uk'],
        'nascimento': (1912, 6, 23),
        'senha': 'enigma',
    },
    {
        'nome': 'Grace Hopper',
        'emails': ['amazing-grace@usnavy.gov'],
        'nascimento': (1906, 12, 9),
        'senha': 'cobol',
    },
    {
        'nome': 'Donald Knuth',
        'emails': ['knuth@gmail.com', 'art-of-programming@caltech.edu'],
        'nascimento': (1938, 1, 10),
        'senha': 'LaTeX',
    },
]

def login(email, senha):
    """
    Retorna um usuário caso o e-mail e senha estejam corretos. Retorna None caso contrário.
    """
    NotImplemented # YOUR ANSWER HERE
    return None


def recuperar_senha(email, nascimento):
    """
    Mostra a senha do usuário ou uma mensagem de erro caso o e-mail seja 
    inválido ou a data de nascimento não corresponda à correta.
    """
    NotImplemented # YOUR ANSWER HERE
    return 'Data de nascimento ou e-mail inválidos'
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de login()"

assert login('turing@gmail.co.uk', 'enigma') == USUARIOS[1]
assert login('turing@gmail.co.uk', 'password-incorreto') is None
assert login('ada@gmail.co.uk', 'algoritmo') == USUARIOS[0]
print('Parabéns!')
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função de login() - múltiplos e-mails"

assert login('knuth@gmail.com', 'LaTeX') == USUARIOS[-1]
assert login('art-of-programming@caltech.edu', 'LaTeX') == USUARIOS[-1]
assert login('art-of-programming@caltech.edu', 'latex') is None
print('Parabéns!')
```

<!-- grade_id: registro-recuperar-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função recuperar_senha() - casos de sucesso"

assert recuperar_senha('turing@gmail.co.uk', (1912, 6, 23)) == 'A senha é "enigma"'
assert recuperar_senha('amazing-grace@usnavy.gov', (1906, 12, 9)) == 'A senha é "cobol"'
print('Parabéns!')
```

<!-- grade_id: registro-recuperar-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função recuperar_senha() - casos de erro"

assert recuperar_senha('turing@gmail.co.uk', (1901, 2, 30)) == 'Data de nascimento ou e-mail inválidos'
assert recuperar_senha('bad@email.com.br', (1912, 6, 23)) == 'Data de nascimento ou e-mail inválidos'
print('Parabéns!')
```