# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Divisão basica

Implemente uma função que retorne a divisão de dois números:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def div(number_1, number_2):
    """
    Retorna a divisão dos dois números informados. Em caso de divisão por 0 deve retornar None.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de div()"

assert div(10, 2) == 5
assert div(55, 0) is None
assert div(-250, 10) == -25
print('Parabéns!')
```