# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Cruzamento de Listas

Dada duas listas de entrada como parâmetro, ex.:

a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

Implemente a função que retorne uma nova lista com os elementos que aparecem em ambas as listas sem repetir números que já foram selecionados.

<!-- grade_id: classe-pessoa, points: 0, grade: false, solution: true -->
```python
def list_overlap(list_a, list_b):
    '''
    Retorna uma lista com os números que aparecem em ambas as listas.
    '''
    NotImplemented # YOUR ANSWER HERE
    return None
```

```python
# Exemplo para você experimentar sua implementação :D
# Deve mostrar a lista: [1, 4]
lista = list_overlap([1, 2, 3, 4], [1, 1, 1, 4])
lista
```

<!-- grade_id: classe-pessoa-test, points: 0, grade: true, solution: false -->
```python
assert list_overlap([1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89], [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]) == [1, 2, 3, 5, 8, 13]
assert list_overlap([0, 1, 2, 4, 5, 50, 56, 100, 101, 200], [0, 2, 3, 5, 40, 45, 46, 56, 100]) == [0, 2, 5 ,56 100]
assert list_overlap([2, 80, 85], [1, 2, 3, 4, 5, 6, 7, 8, 9, 85, 100, 102]) == [2, 85]
print('Parabens')
```