# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Odd Or Even

Implemente uma função que retorne uma String informando se o número inserido é par ou ímpar:

"Esse número é Par" se for par;

"Esse número é impar" se for ímpar;

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def odd_or_even(number):
    """
    Retorna a string da frase correspondente a paridade do número informado.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de odd_or_even()"

assert odd_or_even(1) == "Esse número é impar"
assert odd_or_even(55) == "Esse número é impar"
assert odd_or_even(92) == "Esse número é Par"
assert odd_or_even(852) == "Esse número é Par"
print('Parabéns!')
```