# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Diferença máxima de comprimento

You are given two arrays a1 and a2 of strings. Each string is composed with letters from a to z. Let x be any string in the first array and y be any string in the second array.

Você receberá dois arrays de Strings. Cada String é composta por várias letras de A a Z em quantidades aleatórias. Seu código deve retornar a maior diferença absoluta entre os dois arrays. Ou seja: Find max(abs(length(x) − length(y))) -> Dado que x é uma string do array 1 e y uma string do array 2.

Se uma das listas for vazia o valor retornado deve ser -1.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def mxdiflg(a1, a2):
    """
    Retorna um inteiro com o valor da maior diferença entre as strings dos dois arrays
    """
    NotImplemented
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
# Exemplo para você experimentar sua implementação :)
a1 = ["hoqq", "bbllkw", "oox", "ejjuyyy", "plmiis", "xxxzgpsssa", "xxwwkktt", "znnnnfqknaz", "qqquuhii", "dvvvwz"]
a2 = ["cccooommaaqqoxii", "gggqaffhhh", "tttoowwwmmww"]
print("{}".format(mxdiflg(a1, a2)));       # Imprime o valor da função, que pros arrays a cima deve ser 13
```

Output:
```
13
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
# Testes do Sistema

assert mxdiflg(
    ["hoqq", "bbllkw", "oox", "ejjuyyy", "plmiis", "xxxzgpsssa", "xxwwkktt", "znnnnfqknaz", "qqquuhii", "dvvvwz"],
    ["cccooommaaqqoxii", "gggqaffhhh", "tttoowwwmmww"]
) == 13;
assert mxdiflg(
    ["ejjjjmmtthh", "zxxuueeg", "aanlljrrrxx", "dqqqaaabbb", "oocccffuucccjjjkkkjyyyeehh"],
    ["bbbaaayddqbbrrrv"]
) == 10
assert mxdiflg(
    [],
    ["cccooommaaqqoxii", "gggqaffhhh", "tttoowwwmmww"]
) == -1
assert mxdiflg(
    ["a", "asd", "asdf", "asdfg"],
    ["asdfghjkl", "qwertyuiopasdfgh", "as"]
) == 15
print('Parabens');
```

Output:
```
Parabens
```