# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Verificando números primos

Implemente uma função que informe se um número é primo ou não.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def prime_validation(number):
    """
    Retorna True se o número for primo
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-recuperar-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função prime_validation()"

assert prime_validation(2) == True
assert prime_validation(97) == True
assert prime_validation(1500) == False
assert prime_validation(66) == False
assert prime_validation(83) == True
print('Parabéns!')
```