# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Concatenação básica 

Implemente uma função que retorne a concatenação das duas strings informadas:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def concat(string_1, string_2):
    """
    Retorna uma string concatenada.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de concat()"

assert concat('Shado', 'Frog') == 'ShadoFrog'
assert concat('O', ' Jogo') == 'O Jogo'
assert concat('Shado', 'Gostoso Demais') == 'ShadoGostoso Demais'
print('Perdi!')
```