# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Corrigir número de telefone

Escreva uma função que valida e corrige um número de telefone. Faça um programa que leia um número de telefone, e corrija o número no caso deste conter somente 8 dígitos, acrescentando o '9' na frente. O usuário pode informar o número com ou sem o traço separador. Caso o usuário insira um número de telefone incorrigível com o acréscimo do dígito 9, a função deve retornar a seguinte mensagem: "Telefone inexistente", caso o ele insira um número de telefone válido e correto, a função deve retornar a seguinte mensagem: "Telefone OK".

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def phoneValidate(phoneNumber):
    """
    Retorna o número de telefone correto
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de secondsConverter()"

assert phoneValidate('123') == 'Telefone inexistente.'
assert phoneValidate('-8') == 'Telefone inexistente.'
assert phoneValidate('1234567') == 'Telefone inexistente.'
assert phoneValidate('12345678') == '912345678'
assert phoneValidate('123456789') == 'Telefone OK'
assert phoneValidate('996040414') == 'Telefone OK'
assert phoneValidate('99604-0415') == 'Telefone OK'
assert phoneValidate('8524-2423') == '98524-2423'
assert phoneValidate('85242423') == '985242423'
assert phoneValidate('40028922') == '940028922'
print('Parabéns!')
```