# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Corrigir número de telefone

Faça um programa que peça um numero inteiro positivo e em seguida mostre este número invertido. Ex: 4561 -> 1654. Caso o usuário insira um número inválido, a função deve retornar a seguinte mensagem: "Tente Novamente".

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def numberInverter(number):
    """
    Retorna o número invertido
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de numberInverter()"

assert numberInverter(0) == 'Tente Novamente'
assert numberInverter(-8) == 'Tente Novamente'
assert numberInverter(1234567) == 7654321
assert numberInverter(8954222) == 2224598
assert numberInverter(1) == 1
assert numberInverter(14) == 41
assert numberInverter(5666666) == 6666665
assert numberInverter(753951) == 159357
print('Parabéns!')
```