# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

This question was adapted from this: https://github.com/fabiommendes/desenvolvimento-de-software/blob/master/provas/teste-2.ipynb.
And its use has been authorized by its repository owner.

---
## Multiplição básica

Implemente uma função que realize a multiplicação de 2 números:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def mult(number_1, number_2):
    """
    Retorna o resultado da multiplicação
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de mult()"

assert mult(3, 15) == 45
assert mult(10, -2) == -20
assert mult(-16421, 0) == 0
print('Parabéns!')
```