# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in the place that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Remover Strings duplicadas

Seu código receberá uma lista de strings e deve retornar uma lista com as mesmas strings porém sem repetições caso existam.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(values):
    """
    Retorna uma lista de strings sem duplicatas
    """
    NotImplemented
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
# Exemplo para você experimentar sua implementação :)
values = ["cccooommaaqqoxii", "gggqaffhhh", "tttoowwwmmww", "cccooommaaqqoxii", "tttoowwwmmww"]
print("{}".format(remove_duplicates(values)));       # Imprime o valor da função
```

Output:
```
13
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
# Testes do Sistema

assert remove_duplicates( ["hoqq", "bbllkw", "oox", "oox", "ejjuyyy", "plmiis", "plmiis", "xxxzgpsssa", "xxwwkktt", "znnnnfqknaz", "qqquuhii", "dvvvwz", "dvvvwz", "dvvvwz", "dvvvwz"] ) == ["hoqq", "bbllkw", "oox", "ejjuyyy", "plmiis", "xxxzgpsssa", "xxwwkktt", "znnnnfqknaz", "qqquuhii", "dvvvwz"];
assert remove_duplicates( ["hoqq", "hoqq", "hoqq", "hoqq", "hoqq", "hoqq", "hoqq", "hoqq", "oi", "yes"] ) == ["hoqq", "oi", "yes"];
assert remove_duplicates( ["yes", "test", "test1", "test2", "test3", "test4", "test", "test1", "test2", "test3", "test4", "a"] ) == ["yes", "test", "test1", "test2", "test3", "test4", "a"];
assert remove_duplicates( ["test"] ) == ["test"];
assert remove_duplicates( [] ) == [];
assert remove_duplicates( ["qwer", "qwer", "opa", "opa"] ) == ["qwer", "opa"];
print('Parabens');
```

Output:
```
Parabens
```