# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in all the places that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

This list was adapted from this: https://github.com/fabiommendes/desenvolvimento-de-software/blob/master/provas/teste-1.ipynb.
And its use has been authorized by its repository owner.

---
## Soma de quadrados

Crie uma função que receba um valor inteiro n e retorne a soma dos quadrados de 1 até n. Se n for zero ou negativo, a função deve retornar 0. 

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def soma_dos_quadrados(n):
    """Soma os valores 1 + 2*2 + 3*3 + ... + n*n."""
    NotImplemented # YOUR ASWER HERE
```

A soma dos números de 1 até 10 é 385. Verifique o resultado.

```python
soma_dos_quadrados(10)
```

<!-- grade_id: squares-test, points: 1, grade: true, solution: false -->
```python
"""Verifica se a função foi implementada corretamente."""
assert soma_dos_quadrados(10) == 385
assert soma_dos_quadrados(100) == 338350
assert soma_dos_quadrados(-42) == 0
print('Parabéns!')
```

---
## Fizzbuzz

A função fizzbuzz recebe um número inteiro positivo n e itera sobre a sequência. A cada passo, faz uma das
4 opções:
* Se o número for divisível por 3, imprime "fizz"
* Se o número for divisível por 5, imprime "buzz"
* Se o número for divisível por ambos, imprime "fizzbuzz"
* Se não for divisível por nenhum dos dois, imprime o próprio número.


<!-- grade_id: fizzbuzz, points: 0, grade: false, solution: true -->
```python
def fizzbuzz(n):
    """Imprime a sequência fizzbuzz"""
    NotImplemented # YOUR ANSWER HERE
```

A sequência fizzbuzz de 1 até 10 é dada por
```
1, 2, fizz, 4, buzz, fizz, 7, 8, fizz, buzz, 11, fizz, 13, 14, fizzbuzz, 16, 17, fizz, 19, buzz
```
Você deve imprimir cada valor em uma linha separada.
Verifique sua implementação

```python
fizzbuzz(20)
```

<!-- grade_id: fizzbuzz-short, points: 0.5, grade: true, solution: false -->
```python
"""Verifica o fizzbuzz para valores pequenos"""
import io
from contextlib import redirect_stdout


with redirect_stdout(io.StringIO()) as fd:
    fizzbuzz(10)
    assert  fd.getvalue() == \
"""1
2
fizz
4
buzz
fizz
7
8
fizz
buzz
"""
print('Parabéns!')
```

<!-- grade_id: fizzbuzz-large, points: 0.5, grade: true, solution: false -->
```python
"""Fizzbuzz para valores grandes"""
with redirect_stdout(io.StringIO()) as fd:
    fizzbuzz(20)
    assert  fd.getvalue() == \
"""1
2
fizz
4
buzz
fizz
7
8
fizz
buzz
11
fizz
13
14
fizzbuzz
16
17
fizz
19
buzz
"""
print('Parabéns!')
```

---
## Lista fizzbuzz

Resolva o problema do fizzbuzz, mas desta vez retorne uma lista de strings com os valores de cada iteração ao invés de imprimí-los na tela.

<!-- grade_id: lista-fizzbuzz, points: 0, grade: false, solution: true -->
```python
def lista_fizzbuzz(n):
    """Retorna uma lista com a sequência fizzbuzz"""
    lst = []
    NotImplemented # YOUR ANSWER HERE
```

Teste sua implementação aqui. 
O resultado para `lista_fizzbuzz(20)` é

```python
['1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8', 'fizz', 'buzz', 
 '11', 'fizz', '13', '14', 'fizzbuzz', '16', '17', 'fizz', '19', 'buzz']
```

```python
lista_fizzbuzz(20)
```

<!-- grade_id: lista-fizzbuzz-test, points: 1, grade: true, solution: false -->
```python
"Testes para listas de fizzbuzzes"

assert lista_fizzbuzz(20) == ['1', '2', 'fizz', '4', 'buzz', 'fizz', '7', '8', 'fizz', 'buzz', 
 '11', 'fizz', '13', '14', 'fizzbuzz', '16', '17', 'fizz', '19', 'buzz']
print('Parabéns!')
```

---
## Todos iguais
Crie uma função que verifique se todos os elementos de uma lista são iguais.

Uma lista vazia deve retornar True.

<!-- grade_id: todos-iguais, points: 0, grade: false, solution: true -->
```python
def todos_iguais(lst: list) -> bool:
    NotImplemented # YOUR ANSWER HERE
```

```python
# Verifique sua implementação
todos_iguais([1, 1, 2, 1, 1])
```

<!-- grade_id: todos-iguais-test, points: 1, grade: true, solution: false -->
```python
"Verifica implementação de todos_iguais()"
assert todos_iguais([1, 1, 1]) == True
assert todos_iguais([1, 2, 1]) == False
assert todos_iguais(['a', 'a', 'a']) == True
assert todos_iguais([]) == True
assert todos_iguais([1]) == True
print('Parabéns!')
```

---
## Primos

Um número é primo se for divisível apenas por 1 e por ele mesmo (obs, zero e um não são primos!). Crie uma função que recebe um número e retorne uma lista com todos números menores ou iguais ao número dado.

Lembre-se que para verificar se um número n é primo, basta analisar se o mesmo é divisível por algum número na lista dos primos menores que n.

<!-- grade_id: primos, points: 0, grade: false, solution: true -->
```python
def primos(n: int) -> list:
    NotImplemented # YOUR ANSWER HERE
```

```python
# Teste sua implementação aqui.
# ex.: primos(50) => [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47]
primos(50)
```

<!-- grade_id: primos-test, points: 1.5, grade: true, solution: false -->
```python
"Testes com alguns primos"
assert primos(13) == [2, 3, 5, 7, 11, 13]
assert primos(50) == [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47]
assert primos(1)  == []
print('Parabéns!')
```

---
## Uma questão de classe

A classe Pessoa mostrada abaixo possui os campos de nome e sobrenome.

Implemente a função filho(nome, p1, p2) que produz uma nova pessoa com o sobrenome calculado a partir dos sobrenomes dos pais. No caso, devemos juntar o último sobrenome de p2 com o último sobrenome de p1 (nesta ordem).

<!-- grade_id: classe-pessoa, points: 0, grade: false, solution: true -->
```python
class Pessoa:
    def __init__(self, nome, sobrenome):
        self.nome = nome
        self.sobrenome = sobrenome
        
    # Ensina o Python como imprimir pessoas. Não modifique.
    def __repr__(self):
        return "Pessoa(%r, %r)" % (self.nome, self.sobrenome)
    
    # Ensina o Python como testar se duas pessoas são iguais. Não modifique.
    def __eq__(self, other):
        return self.nome == other.nome and self.sobrenome == other.sobrenome
    
    def ultimo_sobrenome(self):
        """
        Retorna a última palavra do sobrenome da pessoa.
        
        Útil para analizar sobrenomes compostos.
        """
        NotImplemented # YOUR ANSWER HERE
    
def filho(nome: str, p1: Pessoa, p2: Pessoa) -> Pessoa:
    NotImplemented # YOUR ANSWER HERE
```

```python
# Exemplo para você experimentar sua implementação :)
# O filho, neste caso deve ser Pessoa('Sasha', 'Szafir Meneguel')
xuxa = Pessoa("Xuxa", "Meneghel")
pai = Pessoa("Luciano", "Szafir")
filho("Sasha", xuxa, pai)
```

<!-- grade_id: classe-pessoa-test, points: 0, grade: true, solution: false -->
```python
"Testes de 'Uma questão de classe'"
xuxa = Pessoa("Xuxa", "Meneghel")
pai = Pessoa("Luciano", "Szafir")
pele = Pessoa("Edson", "Arantes do Nascimento")

assert filho("Sasha", pai, xuxa) == Pessoa('Sasha', 'Meneghel Szafir')
assert filho("Sasha", xuxa, pai) == Pessoa('Sasha', 'Szafir Meneghel')
assert filho("Edinho", pele, xuxa) == Pessoa('Edinho', 'Meneghel Nascimento')
print('Parabens')
```