# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in all the places that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Média

Implemente uma função que retorne o valor da média dos números informados:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def average(values):
    """
    Retorna o valor da média dos números informados.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de average()"

assert average([7, 5, 21, 12, 566, 65]) == 112.67
assert average([3, 65, 3, 64, 3, 643, 0]) == 111.57
assert average([61, 16, 34, 34, 43, 1, 6]) == 27.86
assert average([852]) == 852
print('Parabéns!')
```

---
## Odd Or Even

Implemente uma função que retorne uma String informando se o número inserido é par ou ímpar:

"Esse número é Par" se for par;

"Esse número é impar" se for ímpar;

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def odd_or_even(number):
    """
    Retorna a string da frase correspondente a paridade do número informado.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de odd_or_even()"

assert odd_or_even(1) == "Esse número é impar"
assert odd_or_even(55) == "Esse número é impar"
assert odd_or_even(92) == "Esse número é Par"
assert odd_or_even(852) == "Esse número é Par"
print('Parabéns!')
```

---
## Concatenação básica 

Implemente uma função que retorne a concatenação das duas strings informadas:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def concat(string_1, string_2):
    """
    Retorna uma string concatenada.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de concat()"

assert concat('Shado', 'Frog') == 'ShadoFrog'
assert concat('O', ' Jogo') == 'O Jogo'
assert concat('Shado', 'Gostoso Demais') == 'ShadoGostoso Demais'
print('Perdi!')
```