# Learning Python

Remember to **restart the kernel** (Kernel$\rightarrow$Restart) and then **execute all cells** (Cell$\rightarrow$Run All) to verify that all solutions are correct. Fill with your code in all the places that say `SUA SOLUÇÃO` or `YOUR ANSWER HERE` and do not change the other cells in this notebook.

---
## Soma Básica

Implemente uma função que retorne o valor da soma de dois números:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def sum(number_1, number_2):
    """
    Retorna o resultado da soma
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de sum()"

assert sum(1, 1) == 2
assert sum(0, 5) == 5
assert sum(2.5, 2.6) == 5.1
assert sum(-5, -20) == -25
print('Parabéns!')
```

---
## Subtração básica

Implemente uma função que realiza a subtração de 2 números: 

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def sub(number_1, number_2):
    """
    Retorna o resultado da subtração de dois números.
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de sub()"

assert sub(1, 1) == 0
assert sub(300, 1560) == -1260
assert sub(-100, -1500) == 1400
print('Parabéns!')
```

---
## Multiplição básica

Implemente uma função que realize a multiplicação de 2 números:

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
def mult(number_1, number_2):
    """
    Retorna o resultado da multiplicação
    """
    NotImplemented # YOUR ANSWER HERE
    return None
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de mult()"

assert mult(3, 15) == 45
assert mult(10, -2) == -20
assert mult(-16421, 0) == 0
print('Parabéns!')
```